<div alert

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name:

```python
NAME = "Hayat Hussein"
```

⚠️⚠️⚠️ **NOTE:** Remember we store your submissions and check against other student's submissions... so, not that you would, but no cheating. ⚠️⚠️⚠️

---

# Project 2: Lists, Iteration, and Strings


## Setting up the autograder

We need to set up the autograder first. Run the cell below just once. This will make sure the autograder works correctly.

(If you need to restart the kernel at any point during your work, make sure you come back to this cell and re-run it again, before using the autograder.)

```python
# This cell needs to be run just once before running the autograder.

from collections import namedtuple
from contextlib import redirect_stdout
from io import StringIO
import re

TEST_PASS_MSG = "Correct, great job! 👍"
TEST_FAIL_MSG = "Something is not quite right!\n" \
    " Check your solution and see if there is anything that needs to be changed."
```

## Pick two problems to solve

Choose ***two*** of the following problems to solve by adding numbers to the markdown cell below. 

Add a '1' to the problem you did for Project 2: Problem Formulation.
Add a '2' to the other problem.

To add numbers, you will need to edit the cell directly; it will not let you click on it. For example, if you want us to grade problem A and that was the one you did the Problem Formulation for, in the cell below you should put an `1` like this:
```
- [1] A
- [ ] B
- [ ] C
```
---

⬇️⬇️⬇️ Enter your response in the cell below ⬇️⬇️⬇️

### Please grade these problems (select two):

- [ ] Problem 1: Tweet Clipper
- [1] Problem 2: Top-secret Word Redacter 
- [2] Problem 3: Email Blocked Domain Filter
- [ ] Problem 4: Rochambo
- [ ] Problem 5: Payment Data Cleaner

---

⚠️⚠️⚠️ You may attempt more than two problems if you want, but if you do, please note that we will count only the two that you mark above. ⚠️⚠️⚠️

## Problem 1: Tweet Clipper


__IMPORTANT__: You can organize your program in any way you want, but you should at least define a function called `tweet_clipper` that takes in input a list of tweets and returns the list of truncated tweets.

### The problem

Tweets can be up to 280 characters long. Write a program that takes in a list of messages of varying length (that are meant to be tweeted) and prints out tweets (one per line) that fits within the 280 characters limit: if the message length is > character limit, truncate it to fit the character limit, and add a "..." at the end of the tweet. (Note: adding the "..." should keep the text within the 280 char limit!) Otherwise, the message remains unchanged. Your program should also print out the number of tweets that had to be truncated because they were too long.

### Example input

For example, given these tweets:
```python
tweets = [
    "Lorem ipsum dolor sit amet, consectetur adipiscing elit. Ut ornare egestas purus quis aliquet. Nulla molestie neque et ante semper, in consequat ante facilisis. Cras sodales, mi et vulputate vestibulum, ante nunc dictum mauris, sed pellentesque sem felis nec sem. Ut feugiat ante magna, pellentesque porttitor lectus sagittis sit amet.",
    "Python is so cool! #programming",
    "Suspendisse nec lacus tempus, sagittis augue non, maximus ante. Cras vehicula, nibh vel sodales tristique, eros libero vestibulum quam, bibendum cursus eros felis ut metus. Suspendisse potenti. Phasellus maximus augue vitae dui luctus, sit amet vestibulum velit lobortis. Class aptent taciti sociosqu ad litora torquent per conubia nostra, per inceptos himenaeos.",
    "Vestibulum lobortis purus at ultrices luctus. Suspendisse hendrerit venenatis mi in posuere. Mauris porta rhoncus nunc in convallis. Cras vel finibus nisl. Suspendisse a euismod elit. Nam congue massa at dolor semper, nec mollis est elementum. Duis accumsan ultrices egestas. Quisque consequat ultricies scelerisque. Vivamus felis turpis, tincidunt sit amet pellentesque eget, malesuada non odio.",
    "It's finally #caturday!"
]
```

your program should produce the following output:

```
TWEET: Lorem ipsum dolor sit amet, consectetur adipiscing elit. Ut ornare egestas purus quis aliquet. Nulla molestie neque et ante semper, in consequat ante facilisis. Cras sodales, mi et vulputate vestibulum, ante nunc dictum mauris, sed pellentesque sem felis nec sem. Ut feugiat an...
TWEET: Python is so cool! #programming
TWEET: Suspendisse nec lacus tempus, sagittis augue non, maximus ante. Cras vehicula, nibh vel sodales tristique, eros libero vestibulum quam, bibendum cursus eros felis ut metus. Suspendisse potenti. Phasellus maximus augue vitae dui luctus, sit amet vestibulum velit lobortis. Class...
TWEET: Vestibulum lobortis purus at ultrices luctus. Suspendisse hendrerit venenatis mi in posuere. Mauris porta rhoncus nunc in convallis. Cras vel finibus nisl. Suspendisse a euismod elit. Nam congue massa at dolor semper, nec mollis est elementum. Duis accumsan ultrices egestas. Q...
TWEET: It's finally #caturday!

Truncated 3 tweets that were over the limit.
```

---

⬇️⬇️⬇️ Enter your solution in the cell below ⬇️⬇️⬇️

<!-- grade_id: cell-1cafbe59eb9d7d3c, points: 0, grade: false, solution: true -->
```python

```

Use the cell below to test your function here (and / or add more code blocks if you need to). 

You should probably play around with these variables to test your function's behavior. 

```python
# example input
test_tweets = [
    "Lorem ipsum dolor sit amet, consectetur adipiscing elit. Ut ornare egestas purus quis aliquet. Nulla molestie neque et ante semper, in consequat ante facilisis. Cras sodales, mi et vulputate vestibulum, ante nunc dictum mauris, sed pellentesque sem felis nec sem. Ut feugiat ante magna, pellentesque porttitor lectus sagittis sit amet.",
    "Python is so cool! #programming",
    "Suspendisse nec lacus tempus, sagittis augue non, maximus ante. Cras vehicula, nibh vel sodales tristique, eros libero vestibulum quam, bibendum cursus eros felis ut metus. Suspendisse potenti. Phasellus maximus augue vitae dui luctus, sit amet vestibulum velit lobortis. Class aptent taciti sociosqu ad litora torquent per conubia nostra, per inceptos himenaeos.",
    "Vestibulum lobortis purus at ultrices luctus. Suspendisse hendrerit venenatis mi in posuere. Mauris porta rhoncus nunc in convallis. Cras vel finibus nisl. Suspendisse a euismod elit. Nam congue massa at dolor semper, nec mollis est elementum. Duis accumsan ultrices egestas. Quisque consequat ultricies scelerisque. Vivamus felis turpis, tincidunt sit amet pellentesque eget, malesuada non odio.",
    "It's finally #caturday!"
]

truncated_tweets = tweet_clipper(test_tweets)
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In[6], line 10
      1 # example input
      2 test_tweets = [
      3     "Lorem ipsum dolor sit amet, consectetur adipiscing elit. Ut ornare egestas purus quis aliquet. Nulla molestie neque et ante semper, in consequat ante facilisis. Cras sodales, mi et vulputate vestibulum, ante nunc dictum mauris, sed pellentesque sem felis nec sem. Ut feugiat ante magna, pellentesque porttitor lectus sagittis sit amet.",
      4     "Python is so cool! #programming",
   (...)
      7     "It's finally #caturday!"
      8 ]
---> 10 truncated_tweets = tweet_clipper(test_tweets)

NameError: name 'tweet_clipper' is not defined
```

When you are done testing your function, run the autograder below to check if your solution is correct. 

The autograder will test that your program:
1. returns the correct truncated tweets; and,
2. prints the correct string `Truncated XX tweets that were over the limit.` (with XX the number of truncated tweets).

<!-- grade_id: cell-b614e7532f73f309, points: 0, grade: true, solution: false -->
```python
pattern = "Truncated (?P<number>\d+) tweets that were over the limit"

TweetClipper = namedtuple("TweetClipper", ["tweets", "expected_tweets", "expected_num_truncated"])

cases = [
    # Test #1
    TweetClipper(
        # Test #1 INPUT
        tweets=[
            "Lorem ipsum dolor sit amet, consectetur adipiscing elit. Ut ornare egestas purus quis aliquet. Nulla molestie neque et ante semper, in consequat ante facilisis. Cras sodales, mi et vulputate vestibulum, ante nunc dictum mauris, sed pellentesque sem felis nec sem. Ut feugiat ante magna, pellentesque porttitor lectus sagittis sit amet.",
            "Python is so cool! #programming",
            "Suspendisse nec lacus tempus, sagittis augue non, maximus ante. Cras vehicula, nibh vel sodales tristique, eros libero vestibulum quam, bibendum cursus eros felis ut metus. Suspendisse potenti. Phasellus maximus augue vitae dui luctus, sit amet vestibulum velit lobortis. Class aptent taciti sociosqu ad litora torquent per conubia nostra, per inceptos himenaeos.",
            "Vestibulum lobortis purus at ultrices luctus. Suspendisse hendrerit venenatis mi in posuere. Mauris porta rhoncus nunc in convallis. Cras vel finibus nisl. Suspendisse a euismod elit. Nam congue massa at dolor semper, nec mollis est elementum. Duis accumsan ultrices egestas. Quisque consequat ultricies scelerisque. Vivamus felis turpis, tincidunt sit amet pellentesque eget, malesuada non odio.",
            "It's finally #caturday!"
        ],
        # Test #1 EXPECTED OUTPUT
        expected_tweets=[
            'Lorem ipsum dolor sit amet, consectetur adipiscing elit. Ut ornare egestas purus quis aliquet. Nulla molestie neque et ante semper, in consequat ante facilisis. Cras sodales, mi et vulputate vestibulum, ante nunc dictum mauris, sed pellentesque sem felis nec sem. Ut feugiat an...',
            'Python is so cool! #programming',
            'Suspendisse nec lacus tempus, sagittis augue non, maximus ante. Cras vehicula, nibh vel sodales tristique, eros libero vestibulum quam, bibendum cursus eros felis ut metus. Suspendisse potenti. Phasellus maximus augue vitae dui luctus, sit amet vestibulum velit lobortis. Class...',
            'Vestibulum lobortis purus at ultrices luctus. Suspendisse hendrerit venenatis mi in posuere. Mauris porta rhoncus nunc in convallis. Cras vel finibus nisl. Suspendisse a euismod elit. Nam congue massa at dolor semper, nec mollis est elementum. Duis accumsan ultrices egestas. Q...',
            "It's finally #caturday!"
        ],
        # Test #1 EXPECTED NUMBER
        expected_num_truncated=3,
    ),
    # Test #2
    TweetClipper(
        # Test #2 INPUT
        tweets=[
            "I got a feeling That tonight's gonna be a good night That tonight's gonna be a good night That tonight's gonna be a good, good night",
            "A feeling That tonight's gonna be a good night That tonight's gonna be a good night That tonight's gonna be a good, good night",
            "A feeling (woo-hoo) That tonight's gonna be a good night That tonight's gonna be a good night That tonight's gonna be a good, good night",
            "A feeling (woo-hoo) That tonight's gonna be a good night That tonight's gonna be a good night That tonight's gonna be a good, good night",
            "Tonight's the night, let's live it up I got my money, let's spend it up (a fee-) Go out and smash it, like, \"Oh my God\" Jump out that sofa, let's kick it, off (a fee-)",
            "I know that we'll have a ball If we get down and go out, and just lose it all I feel, stressed out, I wanna let it go Let's go way out, spaced out, and losin' all control",
            "Fill up my cup, mazel tov Look at her dancin', just take it, off (a fee-) Let's paint the town, we'll shut it down Let's burn the roof, and then we'll do it again",
            "Let's do it, let's do it, let's do it, let's do it And do it, and do it, let's live it up And do it, and do it, and do it, do it, do it Let's do it, let's do it, let's do it",
            "'Cause I gotta feeling (woo-hoo) That tonight's gonna be a good night That tonight's gonna be a good night That tonight's gonna be a good, good night",
            "A feeling (woo-hoo) That tonight's gonna be a good night That tonight's gonna be a good night That tonight's gonna be a good, good night (a fee-)",
            "Tonight's the night (hey), let's live it up (let's live it up) I got my money (I'm paid), let's spend it up (let's spend it up) Go out and smash it (smash it), like, \"Oh my God\" (like, \"Oh my God!\") Jump out that sofa (come on), let's kick it, off (a fee-)",
            "Fill up my cup (drink), mazel tov (l'chaim) Look at her dancing (move it, move it), just take it, off (a fee-) Let's paint the town (paint the town), we'll shut it down (shut it down) Let's burn the roof (ooh-woo), and then we'll do it again",
            "Let's do it, let's do it, let's do it, (let's do it) let's do it And do it (do it), and do it, let's live it up And do it (do it), and do it (and do it) And do it, do it, do it (and do it) Let's do it (and do it), let's do it (and do it) Let's do it (hey), do it (hey), do it (hey), do it",
            "Here we come, here we go, we gotta rock Easy come, easy go, now we on top Feel the shot, body rock, rock it, don't stop Round and round, up and down, around the clock Monday, Tuesday, Wednesday and Thursday (do it) Friday, Saturday, Saturday to Sunday (do it) Get, get, get, get, get with us, you know what we say (say) Party every day, p-p-p-party every day",
            "And I'm feeling (woo-hoo) That tonight's gonna be a good night That tonight's gonna be a good night That tonight's gonna be a good, good night",
            "A feeling (woo-hoo) That tonight's gonna be a good night That tonight's gonna be a good night That tonight's gonna be a good, good night (Woo-hoo)"
        ],
        # Test #2 EXPECTED OUTPUT
        expected_tweets=[
            "I got a feeling That tonight's gonna be a good night That tonight's gonna be a good night That tonight's gonna be a good, good night",
            "A feeling That tonight's gonna be a good night That tonight's gonna be a good night That tonight's gonna be a good, good night",
            "A feeling (woo-hoo) That tonight's gonna be a good night That tonight's gonna be a good night That tonight's gonna be a good, good night",
            "A feeling (woo-hoo) That tonight's gonna be a good night That tonight's gonna be a good night That tonight's gonna be a good, good night",
            'Tonight\'s the night, let\'s live it up I got my money, let\'s spend it up (a fee-) Go out and smash it, like, "Oh my God" Jump out that sofa, let\'s kick it, off (a fee-)',
            "I know that we'll have a ball If we get down and go out, and just lose it all I feel, stressed out, I wanna let it go Let's go way out, spaced out, and losin' all control",
            "Fill up my cup, mazel tov Look at her dancin', just take it, off (a fee-) Let's paint the town, we'll shut it down Let's burn the roof, and then we'll do it again",
            "Let's do it, let's do it, let's do it, let's do it And do it, and do it, let's live it up And do it, and do it, and do it, do it, do it Let's do it, let's do it, let's do it",
            "'Cause I gotta feeling (woo-hoo) That tonight's gonna be a good night That tonight's gonna be a good night That tonight's gonna be a good, good night",
            "A feeling (woo-hoo) That tonight's gonna be a good night That tonight's gonna be a good night That tonight's gonna be a good, good night (a fee-)",
            'Tonight\'s the night (hey), let\'s live it up (let\'s live it up) I got my money (I\'m paid), let\'s spend it up (let\'s spend it up) Go out and smash it (smash it), like, "Oh my God" (like, "Oh my God!") Jump out that sofa (come on), let\'s kick it, off (a fee-)',
            "Fill up my cup (drink), mazel tov (l'chaim) Look at her dancing (move it, move it), just take it, off (a fee-) Let's paint the town (paint the town), we'll shut it down (shut it down) Let's burn the roof (ooh-woo), and then we'll do it again",
            "Let's do it, let's do it, let's do it, (let's do it) let's do it And do it (do it), and do it, let's live it up And do it (do it), and do it (and do it) And do it, do it, do it (and do it) Let's do it (and do it), let's do it (and do it) Let's do it (hey), do it (hey), do it (...",
            "Here we come, here we go, we gotta rock Easy come, easy go, now we on top Feel the shot, body rock, rock it, don't stop Round and round, up and down, around the clock Monday, Tuesday, Wednesday and Thursday (do it) Friday, Saturday, Saturday to Sunday (do it) Get, get, get, ge...",
            "And I'm feeling (woo-hoo) That tonight's gonna be a good night That tonight's gonna be a good night That tonight's gonna be a good, good night",
            "A feeling (woo-hoo) That tonight's gonna be a good night That tonight's gonna be a good night That tonight's gonna be a good, good night (Woo-hoo)"
        ],
        # TEST #2 EXPECTED NUMBER
        expected_num_truncated=2,
    ),
    # Test #3 (source: https://www.gutenberg.org/ebooks/23755)
    TweetClipper(
        # Test #3 INPUT
        tweets=[
            'Himalayan birds inhabit what is perhaps the most wonderful tract of country in the world. The Himalayas are not so much a chain of mountains as a mountainous country, some eighty miles broad and several hundred long--a country composed entirely of mountains and valleys with no large plains or broad plateaux',
            'There is a saying of an ancient Sanskrit poet which, being translated into English, runs: "In a hundred ages of the gods I could not tell you of the glories of Himachal." This every writer on things Himalayan contrives to drag into his composition. Some begin with the quotation, while others reserve it for the last, and make it do duty for the epigram which stylists assure us should terminate every essay',
            'Some there are who quote the Indian sage only to mock him. Such assert that the beauties of the Himalayas have been greatly exaggerated--that, as regards grandeur, their scenery compares unfavourably with that of the Andes, while their beauty is surpassed by that of the Alps. Not having seen the Andes, I am unable to criticise the assertion regarding the grandeur of the Himalayas, but I find it difficult to imagine anything finer than their scenery.'],
        # Test #3 EXPECTED OUTPUT
        expected_tweets=[
            'Himalayan birds inhabit what is perhaps the most wonderful tract of country in the world. The Himalayas are not so much a chain of mountains as a mountainous country, some eighty miles broad and several hundred long--a country composed entirely of mountains and valleys with no...',
            'There is a saying of an ancient Sanskrit poet which, being translated into English, runs: "In a hundred ages of the gods I could not tell you of the glories of Himachal." This every writer on things Himalayan contrives to drag into his composition. Some begin with the quotatio...',
            'Some there are who quote the Indian sage only to mock him. Such assert that the beauties of the Himalayas have been greatly exaggerated--that, as regards grandeur, their scenery compares unfavourably with that of the Andes, while their beauty is surpassed by that of the Alps. ...'],
        # TEST #3 EXPECTED NUMBER
        expected_num_truncated=3,
    ),
    # Test #4 (source: https://www.gutenberg.org/ebooks/25973)
    TweetClipper(
        # Test #4 INPUT
        tweets=[
            'To Study The Birds From The Level Plains To The Crests Of The Peaks Swimming In Cloudland',
            'To Note The Species That Are Peculiar To The Various Altitudes, As Well As Those That Range From The Lower Areas To The Alpine Heights',
            'To Observe The Behavior Of All The Birds Encountered In The West, And Compare Their Habits, Songs, And General Deportment With Those Of Correlated Species And Genera In The East',
            'To Learn As Much As Possible About The Migratory Movements Up And Down The Mountains As The Seasons Wax And Wane,--Surely That Would Be An Inspiring Prospect To Any Student Of The Feathered Fraternity',
            "For Many Years One Of The Writer'S Most Cherished Desires Has Been To Investigate The Bird Life Of The Rocky Mountains",
            'In The Spring Of 1899, And Again In 1901, Fortune Smiled Upon Him In The Most Genial Way, And--In A Mental State Akin To Rapture, It Must Be Confessed--He Found Himself Rambling Over The Plains And Mesas And Through The Deep Cañons, And Clambering Up The Dizzy Heights, In Search Of Winged Rarities'
        ],
        # Test #4 EXPECTED OUTPUT
        expected_tweets=[
            'To Study The Birds From The Level Plains To The Crests Of The Peaks Swimming In Cloudland',
            'To Note The Species That Are Peculiar To The Various Altitudes, As Well As Those That Range From The Lower Areas To The Alpine Heights',
            'To Observe The Behavior Of All The Birds Encountered In The West, And Compare Their Habits, Songs, And General Deportment With Those Of Correlated Species And Genera In The East',
            'To Learn As Much As Possible About The Migratory Movements Up And Down The Mountains As The Seasons Wax And Wane,--Surely That Would Be An Inspiring Prospect To Any Student Of The Feathered Fraternity',
            "For Many Years One Of The Writer'S Most Cherished Desires Has Been To Investigate The Bird Life Of The Rocky Mountains",
            'In The Spring Of 1899, And Again In 1901, Fortune Smiled Upon Him In The Most Genial Way, And--In A Mental State Akin To Rapture, It Must Be Confessed--He Found Himself Rambling Over The Plains And Mesas And Through The Deep Cañons, And Clambering Up The Dizzy Heights, In Sear...'],
        # TEST #4 EXPECTED NUMBER
        expected_num_truncated=1,
    ),
]

# this code takes your function and tests it against the test cases above
print(f"Testing the {tweet_clipper.__name__} function (see code for inputs / expected outputs):")

for i, case in enumerate(cases, start=1):
    print(f"  Test #{i}:")
    
    # run student's function and capture output of any print() in it
    with redirect_stdout(StringIO()) as out:
        tweet_clipper(list(case.tweets))

    # test that program printed the correct output
    print(" A) Prints correct output", end="...")
    try:
        expected_output = '\n'.join(['TWEET: ' + tweet for tweet in case.expected_tweets])
        expected_output += f"\n\nTruncated {case.expected_num_truncated} tweets that were over the limit.\n"
        assert out.getvalue() == expected_output, TEST_FAIL_MSG
        print("OK.")
    except AssertionError:
        print("FAIL!")
        raise
        
# if all test cases passed, print a final message
print()

print(TEST_PASS_MSG)
```

<!-- grade_id: cell-ac9e1333c50334d2, points: 0, grade: true, solution: false -->
```python
# This cell is reserved for grading, please do not change it.

```

## Problem 2: Sensitive word redacter

__IMPORTANT__: You can organize your program in any way you want, but you should at least define a function called `word_redacter` that takes in input the two lists described below and returns the new list of words (in addition to printing it).

### The problem

Write a program that takes a list of words and a list of top-secret words, and prints out a new list of words that replaces each top-secret word with the word "REDACTED", as well as the number of times the top-secret words showed up.


### Example input

For example, given these words and top-secret words:
```python
words = ['she', 'sells', 'sea-shells', 'on', 'the', 'sea-shore', 'the', 'shells', 'she', 'sells', 'are', 'sea-shells', "i'm", 'sure', 'for', 'if', 'she', 'sells', 'sea-shells', 'on', 'the', 'sea-shore', 'then', "i'm", 'sure', 'she', 'sells', 'sea-shore', 'shells']

top_secret_words = ["she", "sea-shells"]

```

your program should produce the following output:

```
REDACTED
sells
REDACTED
on
the
sea-shore
the
shells
REDACTED
sells
are
REDACTED
i'm
sure
for
if
REDACTED
sells
REDACTED
on
the
sea-shore
then
i'm
sure
REDACTED
sells
sea-shore
shells

Redacted 7 words.
```

---

⬇️⬇️⬇️ Enter your solution in the cell below ⬇️⬇️⬇️

<!-- grade_id: cell-af888b82d947e9f9, points: 0, grade: false, solution: true -->
```python
def word_redacter(words, top_secret_words):
    """
    Redacting special words from lists 
    
    parameters 
    words: list
    top_secret_words: list that will be redacted
    
    returns
    returns the list of new words displaying the secret words as "REDACTED"
    """
    redacted_count = 0 #counting redacted words
    
    for word in words: #iteration
        if word in top_secret_words: #if the words are in the top_secret_word list
            print("REDACTED") #prints the word "REDACTED"
            redacted_count += 1 #adds 
        else:
            print(word) #if its not a top_secret_word the actual word prints
    print(f"\nRedacted {redacted_count} {'word' if redacted_count== 1 else 'words'}.") #creates a statement of how many words are redacted
    return(words, top_secret_words) #returns words and top_secret_words that are redacted
```

Use the cell below to test your function here (and / or add more code blocks if you need to). 

You should probably play around with these variables to test your function's behavior. 

```python
test_words = ['she', 'sells', 'sea-shells', 'on', 'the', 'sea-shore', 'the', 'shells', 'she', 'sells', 'are', 'sea-shells', "i'm", 'sure', 'for', 'if', 'she', 'sells', 'sea-shells', 'on', 'the', 'sea-shore', 'then', "i'm", 'sure', 'she', 'sells', 'sea-shore', 'shells']
test_top_secret_words = ['she', 'sea-shells']
redacted_words = word_redacter(test_words, test_top_secret_words)
```

Output:
```
REDACTED
sells
REDACTED
on
the
sea-shore
the
shells
REDACTED
sells
are
REDACTED
i'm
sure
for
if
REDACTED
sells
REDACTED
on
the
sea-shore
then
i'm
sure
REDACTED
sells
sea-shore
shells

Redacted 7 words.
```

When you are done testing your function, run the autograder below to check if your solution is correct. 

The autograder will test that your program:
1. returns the correct redacted words; and,
2. prints the correct string `Redacted XX words.` (with XX the number of redacted words).

<!-- grade_id: cell-7ed1a9dc3647b56b, points: 0, grade: true, solution: false -->
```python
pattern = "Redacted (?P<number>\d+) words"

WordRedacter = namedtuple("WordRedacter", 
                          ["words", 
                           "top_secret_words", 
                           "expected_words", 
                           "expected_num_redacted"])

cases = [
    # Test #1
    WordRedacter(
        # Test #1 INPUT
        words=['she', 'sells', 'sea-shells', 'on', 'the', 'sea-shore', 'the', 'shells', 'she', 'sells', 'are', 'sea-shells', "i'm", 'sure', 'for', 'if', 'she', 'sells', 'sea-shells', 'on', 'the', 'sea-shore', 'then', "i'm", 'sure', 'she', 'sells', 'sea-shore', 'shells'],
        top_secret_words=["she", "sea-shells"],
        # Test #1 EXPECTED OUTPUT
        expected_words=['REDACTED', 'sells', 'REDACTED', 'on', 'the', 'sea-shore', 'the', 'shells', 'REDACTED', 'sells', 'are', 'REDACTED', "i'm", 'sure', 'for', 'if', 'REDACTED', 'sells', 'REDACTED', 'on', 'the', 'sea-shore', 'then', "i'm", 'sure', 'REDACTED', 'sells', 'sea-shore', 'shells'],
        # Test #1 EXPECTED NUMBER
        expected_num_redacted=7,
    ),
    # Test #2
    WordRedacter(
        # Test #2 INPUT
        words=['she', 'sells', 'sea-shells', 'on', 'the', 'sea-shore', 'the', 'shells', 'she', 'sells', 'are', 'sea-shells', "i'm", 'sure', 'for', 'if', 'she', 'sells', 'sea-shells', 'on', 'the', 'sea-shore', 'then', "i'm", 'sure', 'she', 'sells', 'sea-shore', 'shells'],
        top_secret_words=["she", "sea-shells", "sells"],
        # Test #2 EXPECTED OUTPUT
        expected_words=['REDACTED', 'REDACTED', 'REDACTED', 'on', 'the', 'sea-shore', 'the', 'shells', 'REDACTED', 'REDACTED', 'are', 'REDACTED', "i'm", 'sure', 'for', 'if', 'REDACTED', 'REDACTED', 'REDACTED', 'on', 'the', 'sea-shore', 'then', "i'm", 'sure', 'REDACTED', 'REDACTED', 'sea-shore', 'shells'],
        # Test #2 EXPECTED NUMBER
        expected_num_redacted=11,
    ),
    # Test #3
    WordRedacter(
        # Test #3 INPUT
        words=['she', 'sells', 'sea-shells', 'on', 'the', 'sea-shore', 'the', 'shells', 'she', 'sells', 'are', 'sea-shells', "i'm", 'sure', 'for', 'if', 'she', 'sells', 'sea-shells', 'on', 'the', 'sea-shore', 'then', "i'm", 'sure', 'she', 'sells', 'sea-shore', 'shells'],
        top_secret_words=[],
        # Test #3 EXPECTED OUTPUT
        expected_words=['she', 'sells', 'sea-shells', 'on', 'the', 'sea-shore', 'the', 'shells', 'she', 'sells', 'are', 'sea-shells', "i'm", 'sure', 'for', 'if', 'she', 'sells', 'sea-shells', 'on', 'the', 'sea-shore', 'then', "i'm", 'sure', 'she', 'sells', 'sea-shore', 'shells'],
        # Test #3 EXPECTED NUMBER
        expected_num_redacted=0,
    ),
    # Test #4
    WordRedacter(
        # Test #4 INPUT
        words=['mary', 'had', 'a', 'little', 'lamb', 'little', 'lamb', 'little', 'lamb', 'mary', 'had', 'a', 'little', 'lamb', 'its', 'fleece', 'was', 'white', 'as', 'snow'],
        top_secret_words=["lamb"],
        # Test #4 EXPECTED OUTPUT
        expected_words=['mary', 'had', 'a', 'little', 'REDACTED', 'little', 'REDACTED', 'little', 'REDACTED', 'mary', 'had', 'a', 'little', 'REDACTED', 'its', 'fleece', 'was', 'white', 'as', 'snow'],
        # Test #4 EXPECTED NUMBER
        expected_num_redacted=4,
    ),
    # Test #5
    WordRedacter(
        # Test #5 INPUT
        words=[],
        top_secret_words=["lorem", "ipsum"],
        # Test #5 EXPECTED OUTPUT
        expected_words=[],
        # Test #5 EXPECTED NUMBER
        expected_num_redacted=0,
    )
]

# this code takes your function and tests it against the test cases above
print(f"Testing the {word_redacter.__name__} function (see code for inputs / expected outputs):")

for i, case in enumerate(cases, start=1):
    print(f"  Test #{i}:")
    
    # run student's function and capture output of any print() in it
    with redirect_stdout(StringIO()) as out:
        redacted_words = word_redacter(list(case.words), case.top_secret_words)

    # test returned value is same as expecteed
    print("     A) Returns correct redacted words", end="...")
    try:
        assert redacted_words == case.expected_words, TEST_FAIL_MSG
        print("OK.", end=" ")
    except AssertionError:
        print("FAIL!")
        raise
        
    # test that program printed the correct output string
    print(" B) Prints correct string", end="...")
    try:
        m = re.search(pattern, out.getvalue(), re.I)
        assert m is not None, TEST_FAIL_MSG
        assert 'number' in m.groupdict(), TEST_FAIL_MSG
        assert int(m.groupdict()['number']) == case.expected_num_redacted, TEST_FAIL_MSG
        print("OK.")
    except AssertionError:
        print("FAIL!")
        raise
        
# if all test cases passed, print a final message
print()

print(TEST_PASS_MSG)
```

Output:
```
Testing the word_redacter function (see code for inputs / expected outputs):
  Test #1:
     A) Returns correct redacted words...FAIL!
AssertionError                            Traceback (most recent call last)
Cell In[15], line 75
     73 print("     A) Returns correct redacted words", end="...")
     74 try:
---> 75     assert redacted_words == case.expected_words, TEST_FAIL_MSG
     76     print("OK.", end=" ")
     77 except AssertionError:

AssertionError: Something is not quite right!
 Check your solution and see if there is anything that needs to be changed.
```

<!-- grade_id: cell-d7c06ab7a491bbd7, points: 0, grade: true, solution: false -->
```python
# This cell is reserved for grading, please do not change it.

```

## Problem 3: Email blocked domain filter

__IMPORTANT__: You can organize your program in any way you want, but you should at least define a function called `email_blocker` that takes in input the two lists described below and returns the new list of email addresses (in addition to printing it).

### The problem

Write a program that takes a list of emails, and a list of blocked domains (like umd.edu, NOT the whole email), and prints out only &ldquo;cleaned&rdquo; emails (one e-mail per line) that are not from the blocked domains, as well as the number of times any email from a blocked domain showed up.

### Example input

For example, given these emails and blocked domains:
```python
emails = [
  "joelchan@umd.edu",
  "firebrand@gmail.com",
  "funlove@yahoo.com",
  "miro@outlook.com",
  "sp@umd.edu",
  "kjfrog@gmail.com"          
]

domains = ['yahoo.com', 'outlook.com']
```

your program should produce the following output:

```
EMAIL: joelchan@umd.edu
EMAIL: firebrand@gmail.com
EMAIL: sp@umd.edu
EMAIL: kjfrog@gmail.com

Filtered 2 emails from blocked domains.
```

---

⬇️⬇️⬇️ Enter your solution in the cell below ⬇️⬇️⬇️

<!-- grade_id: cell-b5c98bb7508195c8, points: 0, grade: false, solution: true -->
```python
def email_blocker (emails, blocked_domains):
    """
    blocks emails from printing if domain is blocked 
    
    parameters
    emails: list of emails to be filtered
    blocked_domains: list of domains that are blocked 
    
    return
    returns cleaned emails and number of blocked emails
    """
    clean_emails = [] #empty list to store emails
    num_blocked_emails = 0 #sets the count for blocked emails
    
    for email in emails: #iterating through email list
        domain = email.split('@')[-1] #taking each domain from the email
        if domain not in blocked_domains: 
            clean_emails.append(email) #emails are added to clean list if not blocked
        else:
            num_blocked_emails +=1 #count of blocked emails go up one once detected
    
    for clean_email in clean_emails:
        print(f"EMAIL: {clean_email}") #prints list of clean emails 
        
    num_blocked_emails = len(emails)- len(clean_emails) #calculation for num of blocked emails
    print(f"\nFiltered {num_blocked_emails} emails from blocked domains.") #prints number of blocked emails form blocked domains
    return (clean_emails)
    
```

Use the cell below to test your function here (and / or add more code blocks if you need to). 

You should probably play around with these variables to test your function's behavior. 

```python
# example input
test_emails = [
  "joelchan@umd.edu",
  "firebrand@gmail.com",
  "funlove@yahoo.com",
  "miro@outlook.com",
  "sp@umd.edu",
  "kjfrog@gmail.com"          
]

test_domains = ['yahoo.com', 'outlook.com']

cleaned_emails = email_blocker(test_emails, test_domains)
```

When you are done testing your function, run the autograder below to check if your solution is correct. 

The autograder will test that your program:
1. returns the correct cleaned emails; and,
2. prints the correct string `Filtered XX emails from blocked domains.` (with XX the number of filtered addresses).

<!-- grade_id: cell-d63cb74b61a26b55, points: 0, grade: true, solution: false -->
```python
pattern = "Filtered (?P<number>\d+) emails from blocked domains"

EmailBlocker = namedtuple("EmailBlocker", ["emails", "domains", "expected_emails", "expected_num_blocked"])

cases = [
    # Test #1
    EmailBlocker(
        # Test #1 INPUT
        emails=["joelchan@umd.edu", "firebrand@gmail.com", "funlove@yahoo.com", "miro@outlook.com", "sp@umd.edu", "kjfrog@gmail.com"],
        domains=['yahoo.com', 'outlook.com'],
        # Test #1 EXPECTED OUTPUT
        expected_emails=['joelchan@umd.edu', 'firebrand@gmail.com', 'sp@umd.edu', 'kjfrog@gmail.com'],
        # Test #1 EXPECTED NUMBER
        expected_num_blocked=2,
    ),
    # Test #2
    EmailBlocker(
        # Test #2 INPUT
        emails=["joelchan@umd.edu", "firebrand@gmail.com", "funlove@yahoo.com", "miro@outlook.com", "sp@umd.edu", "kjfrog@gmail.com"],
        domains=['yahoo.com'],
        # Test #2 EXPECTED OUTPUT
        expected_emails=['joelchan@umd.edu', 'firebrand@gmail.com', 'miro@outlook.com', 'sp@umd.edu', 'kjfrog@gmail.com'],
        # Test #2 EXPECTED NUMBER
        expected_num_blocked=1,
    ),
    # Test #3
    EmailBlocker(
        # Test #3 INPUT
        emails=["joelchan@umd.edu", "firebrand@gmail.com", "funlove@yahoo.com", "miro@outlook.com", "sp@umd.edu", "kjfrog@gmail.com"],
        domains=['yahoo.com', 'outlook.com', "umd.edu"],
        # Test #3 EXPECTED OUTPUT
        expected_emails=['firebrand@gmail.com', 'kjfrog@gmail.com'],
        # Test #3 EXPECTED NUMBER
        expected_num_blocked=4,
    ),
    # Test #4
    EmailBlocker(
        # Test #4 INPUT
        emails=[],
        domains=['yahoo.com', 'outlook.com'],
        # Test #4 EXPECTED OUTPUT
        expected_emails=[],
        # Test #4 EXPECTED NUMBER
        expected_num_blocked=0,
    ),
    # Test #5
    EmailBlocker(
        # Test #5 INPUT
        emails=["joelchan@umd.edu", "firebrand@gmail.com", "funlove@yahoo.com", "miro@outlook.com", "sp@umd.edu", "kjfrog@gmail.com"],
        domains=[],
        # Test #5 EXPECTED OUTPUT
        expected_emails=["joelchan@umd.edu", "firebrand@gmail.com", "funlove@yahoo.com", "miro@outlook.com", "sp@umd.edu", "kjfrog@gmail.com"],
        # Test #5 EXPECTED NUMBER
        expected_num_blocked=0,
    ),
]

# this code takes your function and tests it against the test cases above
print(f"Testing the {email_blocker.__name__} function (see code for inputs / expected outputs):")

for i, case in enumerate(cases, start=1):
    print(f"  Test #{i}:")
    
    # run student's function and capture output of any print() in it
    with redirect_stdout(StringIO()) as out:
        cleaned_emails = email_blocker(list(case.emails), case.domains)

    # test returned value is same as expecteed
    print("     A) Returns correct cleaned emails", end="...")
    try:
        assert cleaned_emails == case.expected_emails, TEST_FAIL_MSG
        print("OK.", end=" ")
    except AssertionError:
        print("FAIL!")
        raise
        
    # test that program printed the correct output string
    print(" B) Prints correct string", end="...")
    try:
        m = re.search(pattern, out.getvalue(), re.I)
        assert m is not None, TEST_FAIL_MSG
        assert 'number' in m.groupdict(), TEST_FAIL_MSG
        assert int(m.groupdict()['number']) == case.expected_num_blocked, TEST_FAIL_MSG
        print("OK.")
    except AssertionError:
        print("FAIL!")
        raise
        
# if all test cases passed, print a final message
print()

print(TEST_PASS_MSG)
```

<!-- grade_id: cell-085527712278a7f3, points: 0, grade: true, solution: false -->
```python
# This cell is reserved for grading, please do not change it.

```

## Problem 4: Rochambo 

__IMPORTANT__: You can organize your program in any way you want, but you should at least define a function called `rochambo` that takes in input the two lists described below and return an integer encoding the outcome of the game (`0` &ndash; draw, `1` &ndash; player 1 won, `2` &ndash; player 2 won), in addition to printing the messages below.

### The problem

Write a program that plays Rochambo (also known as &ldquo;Rock-Paper-Scissors&rdquo;). The program takes two lists of entries, of any length (as long as they are the same length), and computes how many wins each player gets in total, then returns a message that prints out `Player X wins, with N wins to Y wins for Player Z` if one player wins. If it comes to a draw, print out `Draw! Each player has N wins apiece`.

### Example input

For example, given these entries:
```python
player_1 = ["R", "S", "P", "S", "S", "R", "P"]
player_2 = ["S", "S", "R", "R", "P", "S", "P"]
```

your program should print the following output:

```
Player 1 wins, with 4 wins to 1 wins for Player 2
```

And return `1`.

---

⬇️⬇️⬇️ Enter your solution in the cell below ⬇️⬇️⬇️

<!-- grade_id: cell-09ba3c1c2db9aaa4, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Use the cell below to test your function here (and / or add more code blocks if you need to). 

You should probably play around with these variables to test your function's behavior. 

```python
# example inputs
test_player_1 = ["R", "S", "P", "S", "S", "R", "P"]
test_player_2 = ["S", "S", "R", "R", "P", "S", "P"]

return_code = rochambo(test_player_1, test_player_2)
```

When you are done testing your function, run the autograder below to check if your solution is correct. 

The autograder will test that your program:
1. returns the correct code (`0`, `1`, or `2`); and,
2. prints the correct output, either 
  - `Player X wins, with N wins to Y wins for Player Z.`, or 
  - `Draw! Each player has N wins apiece` 

depending on the outcome of the game.

<!-- grade_id: cell-5056118c1341cc39, points: 0, grade: true, solution: false -->
```python
pattern1 = "Player (?P<winner>[12]) wins, with (?P<wins_winner>\d+) wins to (?P<wins_loser>\d+) wins for Player (?P<loser>[12])"
pattern2 = "Draw! Each player has (?P<wins_both>\d+) wins apiece"
pattern = pattern1 + "|" + pattern2

Rochambo = namedtuple("Rochambo",
                      [
                          "player1", 
                          "player2", 
                          "exp_code", 
                          "exp_winner",
                          "exp_wins_winner",
                          "exp_wins_loser",
                          "exp_loser",
                          "exp_wins_both"
                      ])

cases = [
    # Test #1
    Rochambo(
        # Test #1 INPUT
        player1=["R", "S", "P", "S", "S", "R", "P"],
        player2=["S", "S", "R", "R", "P", "S", "P"],
        # Test #1 EXPECTED OUTPUT
        exp_code=1,
        # Test #1 EXPECTED NUMBERS
        exp_winner="1",
        exp_wins_winner="4",
        exp_wins_loser="1",
        exp_loser="2",
        exp_wins_both=None
    ),
    # Test #2
    Rochambo(
        # Test #2 INPUT
        player1=["S", "S", "R", "R", "P", "S", "P"],
        player2=["R", "S", "P", "S", "S", "R", "P"],
        # Test #2 EXPECTED OUTPUT
        exp_code=2,
        # Test #2 EXPECTED NUMBERS
        exp_winner="2",
        exp_wins_winner="4",
        exp_wins_loser="1",
        exp_loser="1",
        exp_wins_both=None
    ),
    # Test #3
    Rochambo(
        # Test #3 INPUT
        player1=["R", "S", "P", "S"],
        player2=["S", "S", "R", "R"],
        # Test #3 EXPECTED OUTPUT
        exp_code=1,
        # Test #3 EXPECTED NUMBERS
        exp_winner="1",
        exp_wins_winner="2",
        exp_wins_loser="1",
        exp_loser="2",
        exp_wins_both=None
    ),
    # Test #4
    Rochambo(
        # Test #4 INPUT
        player1=["S", "P", "S"],
        player2=["S", "R", "R"],
        # Test #4 EXPECTED OUTPUT
        exp_code=0,
        # Test #4 EXPECTED NUMBERS
        exp_winner=None,
        exp_wins_winner=None,
        exp_wins_loser=None,
        exp_loser=None,
        exp_wins_both="1"
    ),
    # Test #5
    Rochambo(
        # Test #5 INPUT
        player1=[],
        player2=[],
        # Test #5 EXPECTED OUTPUT
        exp_code=0,
        # Test #5 EXPECTED NUMBERS
        exp_winner=None,
        exp_wins_winner=None,
        exp_wins_loser=None,
        exp_loser=None,
        exp_wins_both="0"
    ),

]

# this code takes your function and tests it against the test cases above
print(f"Testing the {rochambo.__name__} function (see code for inputs / expected outputs):")

for i, case in enumerate(cases, start=1):
    print(f"  Test #{i}:")
    
    # run student's function and capture output of any print() in it
    with redirect_stdout(StringIO()) as out:
        return_code = rochambo(case.player1, case.player2)

    # test returned value is same as expecteed
    print("     A) Returns correct result code (0/1/2)", end="...")
    try:
        assert return_code == case.exp_code, TEST_FAIL_MSG
        print("OK.", end=" ")
    except AssertionError:
        print("FAIL!")
        raise
        
    # test that program printed the correct output string
    print(" B) Prints correct string", end="...")
    try:
        m = re.search(pattern, out.getvalue(), re.I)
        assert m is not None, TEST_FAIL_MSG
        assert 'winner' in m.groupdict(), TEST_FAIL_MSG
        assert 'loser' in m.groupdict(), TEST_FAIL_MSG
        assert 'wins_winner' in m.groupdict(), TEST_FAIL_MSG
        assert 'wins_loser' in m.groupdict(), TEST_FAIL_MSG
        assert 'wins_both' in m.groupdict(), TEST_FAIL_MSG
        assert m.groupdict()['winner'] == case.exp_winner, TEST_FAIL_MSG
        assert m.groupdict()['loser'] == case.exp_loser, TEST_FAIL_MSG
        assert m.groupdict()['wins_winner'] == case.exp_wins_winner, TEST_FAIL_MSG
        assert m.groupdict()['wins_loser'] == case.exp_wins_loser, TEST_FAIL_MSG
        assert m.groupdict()['wins_both'] == case.exp_wins_both, TEST_FAIL_MSG
        print("OK.")
    except AssertionError:
        print("FAIL!")
        raise
        
# if all test cases passed, print a final message
print()

print(TEST_PASS_MSG)
```

<!-- grade_id: cell-7efd8cb07ddd6a72, points: 0, grade: true, solution: false -->
```python
# This cell is reserved for grading, please do not change it.

```

## Problem 5: Data cleaner

__IMPORTANT__: You can organize your program in any way you want, but you should at least define a function called `data_cleaner` that takes in input the list described below and returns the list of cleaned data (in addition to printing the average).

### The problem

Write a program that takes in a list of payment records that are &ldquo;dirty&rdquo;: they need to be screened out of &ldquo;outliers&rdquo; (anything below \\$0 or above \\$10,000 is an outlier), and then prints out the average payment values for the list. Also note that the payment records are strings (e.g., `"$100,000"`) that should be represented as numbers.

### Example input

For example, given these entries:
```python
payments = [
   "$1,000,000",
   "$-90",
   "$50",
   "$153",
   "$22",
   "$265",
   "$500,000"
]
```

your program should print the following output:

```
Average payment = $122.50
```

with rounding to 2 decimal digits.

---

⬇️⬇️⬇️ Enter your solution in the cell below ⬇️⬇️⬇️

<!-- grade_id: cell-7cd29df659fd5ef3, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Use the cell below to test your function here (and / or add more code blocks if you need to). 

You should probably play around with these variables to test your function's behavior. 

```python
# example input
test_payments = [
   "$1,000,000",
   "$-90",
   "$50",
   "$153",
   "$22",
   "$265",
   "$500,000"
]

cleaned_payments = data_cleaner(test_payments)
```

When you are done testing your function, run the autograder below to check if your solution is correct. 

The autograder will test that your program:
1. returns the correct cleaned data; and,
2. prints the correct string `Average payment = $XX` (with XX the average payment rounded to two decimal digits).

<!-- grade_id: cell-48eaf4fd5d43d3f5, points: 0, grade: true, solution: false -->
```python
pattern = "Average payment = \$(?P<number>\d+.\d\d)"

DataCleaner = namedtuple("DataCleaner", ["payments", "expected_payments", "expected_num_average"])

cases = [
    # Test #1
    DataCleaner(
        # Test #1 INPUT
        payments=["$1,000,000", "$-90", "$50", "$153", "$22", "$265", "$500,000"],
        # Test #1 EXPECTED OUTPUT
        expected_payments=[50, 153, 22, 265],
        # Test #1 EXPECTED NUMBER
        expected_num_average="122.50",
    ),
    # Test #2
    DataCleaner(
        # Test #2 INPUT
        payments=["$1,000", "$2,000", "$-10"],
        # Test #2 EXPECTED OUTPUT
        expected_payments=[1000, 2000],
        # Test #2 EXPECTED NUMBER
        expected_num_average="1500.00",
    ),
    # Test #3
    DataCleaner(
        # Test #3 INPUT
        payments=["$10,000", "$10,000", "$10,000"],
        # Test #3 EXPECTED OUTPUT
        expected_payments=[10000, 10000, 10000],
        # Test #3 EXPECTED NUMBER
        expected_num_average="10000.00",
    ),
    # Test #4
    DataCleaner(
        # Test #4 INPUT
        payments=["$0", "$1", "$11"],
        # Test #4 EXPECTED OUTPUT
        expected_payments=[0, 1, 11],
        # Test #4 EXPECTED NUMBER
        expected_num_average="4.00",
    ),
    # Test #5
    DataCleaner(
        # Test #5 INPUT
        payments=["$1", "$-10", "$10,001"],
        # Test #5 EXPECTED OUTPUT
        expected_payments=[1],
        # Test #5 EXPECTED NUMBER
        expected_num_average="1.00",
    ),
]

# this code takes your function and tests it against the test cases above
print(f"Testing the {data_cleaner.__name__} function (see code for inputs / expected outputs):")

for i, case in enumerate(cases, start=1):
    print(f"  Test #{i}:")
    
    # run student's function and capture output of any print() in it
    with redirect_stdout(StringIO()) as out:
        cleaned_data = data_cleaner(case.payments)

    # test returned value is same as expecteed
    print("     A) Returns correct cleaned emails", end="...")
    try:
        assert cleaned_data == case.expected_payments, TEST_FAIL_MSG
        print("OK.", end=" ")
    except AssertionError:
        print("FAIL!")
        raise
        
    # test that program printed the correct output string
    print(" B) Prints correct string", end="...")
    try:
        m = re.search(pattern, out.getvalue(), re.I)
        assert m is not None, TEST_FAIL_MSG
        assert 'number' in m.groupdict(), TEST_FAIL_MSG
        assert m.groupdict()['number'] == case.expected_num_average, TEST_FAIL_MSG
        print("OK.")
    except AssertionError:
        print("FAIL!")
        raise
        
# if all test cases passed, print a final message
print()

print(TEST_PASS_MSG)
```

<!-- grade_id: cell-13efcffb6e1cb2dc, points: 0, grade: true, solution: false -->
```python
# This cell is reserved for grading, please do not change it.

```