<div alert

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name below:

```python
NAME = "Hayat Hussein"
```

⚠️⚠️⚠️ **NOTE:** Remember we store your submissions and check against other student's submissions... so, not that you would, but no cheating. ⚠️⚠️⚠️

---

# Project 1: Expressions, Variables, Functions, and Conditionals


## Metacognition questions

_Complete after you are finished, just before submitting._

**Q1. While/after completing this project, what (if any) concepts/skills did you realize you feel you need/want to understand better?**

While completing this project, the concepts that were a bit difficult for me to grasp is using operaters in an if, elif, and else statement. Sometimes I had difficulty understanding which operator to use for each problem. I also want to understand the difference between nested conditionals and chained conditionals. While working on this project, I didn't understand if my work exemplified nested conditionals or chained conditionals. 

---

**Q2. How much have you relied on "trial and error"^ to complete this project? Give a rough, ballpark estimate. Answers might range from "basically none" to "almost all of the code was written through trial and error".**

Almost all the code was written through trial and error. I had to rewrite code multiple times for it to run, and even when it ran, I did not write the correct answer. Revisiting my mistakes I rewrote code nearly 20 times before the autograder stated I completed it correctly. 
<br/>
<br/>
^ "trial and error" roughly amounts to blindly changing things until they work, with no predictions about whether your changes are likely to be successful or not.

---

## Setting up the autograder

We need to set up the autograder first. Run the cell below just once. This will make sure the autograder works correctly.

(If you need to restart the kernel at any point during your work, make sure you come back to this cell and re-run it again, before using the autograder.)

```python
# This cell needs to be run just once before running the autograder.

from collections import namedtuple

TEST_PASS_MSG = "Correct, great job! 👍"
TEST_FAIL_MSG = "Something is not quite right!\n" \
    " Check your solution and see if there is anything that needs to be changed."
```

---

## Pick two problems to solve

Choose ***two*** of the following problems to solve by adding checkmarks to the markdown cell below. 

To add checkmarks, you will need to edit the cell directly; it will not let you click on it. For example, if you want us to grade problem A, in the cell below you should put an `X` into the markdown like this and run the cell:
```
- [X] A
- [ ] B
- [ ] C
```

### Please grade this problem (check two):

- [ ] Problem 1: Wedding seating
- [X] Problem 2: COVID testing
- [ ] Problem 3: Football touchdown rules
- [ ] Problem 4: Warrior attack effect calculator
- [X] Problem 5: Coffee recommender

---

⚠️⚠️⚠️ You may attempt more than two problems if you want, but if you do, please note that we will count only the two that you mark here. ⚠️⚠️⚠️

**NOTE:** You will be graded on how well you answer the problems (i.e., arriving at the correct solution) **AND** how you answer the problems (i.e., approaching the problems correctly). For the latter, this means structuring and commenting your code (#'s and docstrings) and using naming conventions according to the practices taught so far in the course. 

## Problem 1: Wedding seating

Write a function called `wedding_seating` that tells you where to sit (*front-row*, *left*, or *right*) depending on these parameters: 
- `relationship`: whether you're **friend** or **family**
- `close_to` whether you're closest to the **bride**, **groom**, or **both**

Here is the mapping of cases to decisions that you want to follow:
- **family** &rarr; *front-row* (no matter if close to bride, groom, or both)
- **friend** of **bride** &rarr; *right*
- **friend** of **groom** &rarr; *left*
- **friend** of **both** bride *and* groom &rarr; *front-row*

<!-- grade_id: cell-e59a5396e0fed9c5, points: 0, grade: false, solution: true -->
```python

```

Use the cell below to test your function here (and / or add more code blocks if you need to). 

You should probably play around with these variables to test your function's behavior. 

```python
relationship = "family"
close_to = "both"

result = wedding_seating(relationship, close_to)
print(result)
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In[3], line 4
      1 relationship = "family"
      2 close_to = "both"
----> 4 result = wedding_seating(relationship, close_to)
      5 print(result)

NameError: name 'wedding_seating' is not defined
```

When you are done testing your function, run the autograder below to check if your solution is correct. 

```python
WeddingCase = namedtuple("WeddingCase", ["relationship", "close_to", "expected_location"])

# the combination of inputs that will be used to test your function
cases = [
    WeddingCase(relationship="family", close_to="groom", expected_location="front-row"),
    WeddingCase(relationship="family", close_to="bride", expected_location="front-row"),
    WeddingCase(relationship="friend", close_to="groom", expected_location="left"),
    WeddingCase(relationship="friend", close_to="bride", expected_location="right"),
    WeddingCase(relationship="friend", close_to="both", expected_location="front-row"),
    WeddingCase(relationship="family", close_to="both", expected_location="front-row")
    
]

# this code takes your function and tests it against the test cases above
print(f"Testing the {wedding_seating.__name__} function:")
for case in cases:
    print(f"  With inputs: relationship={case.relationship!r}, close_to={case.close_to!r}, it "\
          f"should return: {case.expected_location!r}", end="...")
    
    # run student's function
    seat_location = wedding_seating(relationship=case.relationship, close_to=case.close_to)

    # test returned value is same as expecteed
    try:
        assert seat_location == case.expected_location, TEST_FAIL_MSG
        print("OK")
    except AssertionError:
        print("FAIL")
        raise

# if all test cases passed, print a final message
print(TEST_PASS_MSG)
```

<!-- grade_id: cell-8c357ccbcad44c10, points: 0, grade: true, solution: false -->
```python
# This cell is reserved for grading, please do not change it.

```

---

## Problem 2: COVID test
Write a function called `checkforcovid` that makes a tentative diagnosis of COVID based on a set of symptoms. 

The symptoms (parameters) are:
- `temperature` (int) - what is your temperature?
- `fatigue` (Boolean) - are you tired?
- `runny_nose` (Boolean) - do you have a runny nose?

The conditions and outcomes are:
- runny nose is an irrelevant symptom (a [red herring](https://en.wikipedia.org/wiki/Red_herring) if you will), so any time `runny_nose` is true, the function should diagnose as `probably_not` (i.e., not COVID)
- if either temperature is above 100 (fever) or you are tired (but not both!), the function should diagnose as `possibly_yes`
- if temperature is above 100 (fever) and you also tired, the function should diagnose as `probably_yes`

*NOTE: it should go without saying, but these conditions are made up for this exercise: the reality is far more complex!*

<!-- grade_id: cell-d8efddb85a369439, points: 0, grade: false, solution: true -->
```python
#testing for covid based on symptoms and temp
def checkforcovid(temperature, fatigue, runny_nose):
    """
    this function checks for covid test results based on symptoms
    
    Parameters 
    temperature: int
    fatigue: bool
    runny_nose: bool
    
    Returns:
    return --> covid test results : possibly yes, probably yes, probably not
    """
    if runny_nose == True:
        return "probably_not"
    elif temperature >100 and fatigue == True:
        return "probably_yes"
    elif temperature > 100 or fatigue == True:
        return "possibly_yes"
    else:
        return "probably_not"
```

Use the cell below to test your function here (and / or add more code blocks if you need to). 

You should probably play around with these variables to test your function's behavior. 

```python
runny_nose = True
temperature = 101
fatigue = False
result = checkforcovid(temperature, fatigue, runny_nose)
print(result)
```

Output:
```
probably_not
```

When you are done testing your function, run the autograder below to check if your solution is correct. 

```python
DiagnosisCase = namedtuple("DiagnosisCase", ["temperature", "fatigue", "runny_nose", "exp_diagnosis"])

cases = [
    DiagnosisCase(temperature=96, fatigue=False, runny_nose=False, exp_diagnosis="probably_not"),
    DiagnosisCase(temperature=97, fatigue=False, runny_nose=True, exp_diagnosis="probably_not"),
    DiagnosisCase(temperature=98, fatigue=True, runny_nose=False, exp_diagnosis="possibly_yes"),
    DiagnosisCase(temperature=99, fatigue=True, runny_nose=True, exp_diagnosis="probably_not"),
    DiagnosisCase(temperature=101, fatigue=False, runny_nose=False, exp_diagnosis="possibly_yes"),
    DiagnosisCase(temperature=102, fatigue=False, runny_nose=True, exp_diagnosis="probably_not"),
    DiagnosisCase(temperature=103, fatigue=True, runny_nose=False, exp_diagnosis="probably_yes"),
    DiagnosisCase(temperature=104, fatigue=True, runny_nose=True, exp_diagnosis="probably_not")
]

print(f"Testing the {checkforcovid.__name__} function:")
for case in cases:
    print(f"  With inputs: temperature={case.temperature!r}, fatigue={case.fatigue!r}, "\
          f"runny_nose={case.runny_nose!r}, it should return: {case.exp_diagnosis!r}. "\
          f"Outcome", end="...")
    
    # run student's function
    diagnosis = checkforcovid(temperature=case.temperature, fatigue=case.fatigue, runny_nose=case.runny_nose)

    # test returned value is same as expecteed
    try:
        assert diagnosis == case.exp_diagnosis, TEST_FAIL_MSG
        print("OK")
    except AssertionError:
        print("FAIL")
        raise

# if all test cases passed, print final message
print(TEST_PASS_MSG)
```

Output:
```
Testing the checkforcovid function:
  With inputs: temperature=96, fatigue=False, runny_nose=False, it should return: 'probably_not'. Outcome...OK
  With inputs: temperature=97, fatigue=False, runny_nose=True, it should return: 'probably_not'. Outcome...OK
  With inputs: temperature=98, fatigue=True, runny_nose=False, it should return: 'possibly_yes'. Outcome...OK
  With inputs: temperature=99, fatigue=True, runny_nose=True, it should return: 'probably_not'. Outcome...OK
  With inputs: temperature=101, fatigue=False, runny_nose=False, it should return: 'possibly_yes'. Outcome...OK
  With inputs: temperature=102, fatigue=False, runny_nose=True, it should return: 'probably_not'. Outcome...OK
  With inputs: temperature=103, fatigue=True, runny_nose=False, it should return: 'probably_yes'. Outcome...OK
  With inputs: temperature=104, fatigue=True, runny_nose=True, it should return: 'probably_not'. Outcome...OK
Correct, great job! 👍
```

<!-- grade_id: cell-5bab9c6aabb43159, points: 0, grade: true, solution: false -->
```python
# This cell is reserved for grading, please do not change it.

```

---

## Problem 3: Football touchdown rules

Write a function called `touchdown_rules` that calls a touchdown based on rules of football. 

Per [this simplified guide](https://www.thesun.co.uk/sport/5192047/nfl-rules-beginners-guide-touchdown-first-down-quarterback-endzone-super-bowl/#:~:text=A%20player%20can%20run%20the,in%20control%20from%20the%20receiver.), you can score a touchdown by running the ball in or catching a pass in the endzone. If you're running, the ball has only has to "break the plane" (i.e., cross the endzone line). But if you're catching, there are more requirements: both feet have to touch the ground AND the ball must be "in control" of the receiver; if either of these conditions is not met, then it's not a touchdown.

In other words, your function needs to take in these parameters:
- `break_plane` (Boolean) - did the ball "break the plane" (i.e., enter the end zone)?
- `manner` (str) - how did the ball enter the end zone? "run" or "catch"?
- `both_feet` (Boolean) - were both feet on the ground at the time of the call?
- `in_control` (Boolean) - did the player "have control" of the ball at the time of the call?

And return either `True` (touchdown!) or `False` (not touchdown), following the logic laid out above.

<!-- grade_id: cell-c198b940902dc47c, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Test your function here (and / or add more code blocks if you need to). 

You should probably play around with these variables to test your function's behavior.

```python
break_plane = True
manner = "run"
both_feet = True
in_control = True
call = touchdown_rules(break_plane, manner, both_feet, in_control)
print(call)
```

When you are done testing your function, run the autograder below to check if your solution is correct. 

The autograder will test several possible combinations of inputs at once.

<!-- grade_id: cell-9e135fa9d6376255, points: 0, grade: true, solution: false -->
```python
TouchdownCase = namedtuple("TouchdownCase", ["break_plane", "manner", "both_feet", "in_control", "expected_call"])

cases = [
    TouchdownCase(break_plane=False, manner="run", both_feet=True, in_control=False, expected_call=False),
    TouchdownCase(break_plane=True, manner="run", both_feet=False, in_control=True, expected_call=True),
    TouchdownCase(break_plane=False, manner="catch", both_feet=False, in_control=True, expected_call=False),
    TouchdownCase(break_plane=False, manner="catch", both_feet=True, in_control=True, expected_call=True),
    TouchdownCase(break_plane=True, manner="catch", both_feet=True, in_control=False, expected_call=False),
]


print(f"Testing the {touchdown_rules.__name__} function:")
for case in cases:
    print(f"  With: break_plan={case.break_plane!r}, manner={case.manner!r}, "\
          f"both_feet={case.both_feet!r}, in_control={case.in_control!r} it should "\
          f"return: {case.expected_call!r}. Outcome", end="...")
    
    # run student's function
    call = touchdown_rules(break_plane=case.break_plane, manner=case.manner, 
                         both_feet=case.both_feet, in_control=case.in_control)

    # test returned value is same as expecteed
    try:
        assert call == case.expected_call, TEST_FAIL_MSG
        print("OK")
    except AssertionError:
        print("FAIL")
        raise

# if all test cases passed, print final message
print(TEST_PASS_MSG)
```

<!-- grade_id: cell-dfe93c328f6a81b3, points: 0, grade: true, solution: false -->
```python
# This cell is reserved for grading, please do not change it.

```

---

## Problem 4: Warrior attack effect calculator

Write a function called `attack_effect` that calculates whether a player's attack on an opponent was effective or not, depending on the following parameters:
- `attack` (str) - "normal" or "special" attack
- `magic` (Boolean) - does the attack use magic?
- `enemy_hero` (Boolean) - is the enemy a "hero"?

The conditions and outcomes are:
- If your enemy is not a hero, and you use magic and a special attack, the attack is `super_effective`
- If your enemy is not a hero, and you don't use magic or only use a normal attack, the attack is `effective`
- If your enemy is a hero, and you use special magic, the attack is `effective`
- If your enemy is a hero, and you don't use magic, the attack is `ineffective`


<!-- grade_id: cell-2df215a8e1415620, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Test your function here (and / or add more code blocks if you need to). 

You should probably play around with these variables to test your function's behavior.

```python
attack = "special"
magic = False
enemy_hero = False
result = attack_effect(attack, magic, enemy_hero)
print(result)
```

When you are done testing your function, run the autograder below to check if your solution is correct. 

The autograder will test several possible combinations of inputs at once.

<!-- grade_id: cell-69874f6af587bde3, points: 0, grade: true, solution: false -->
```python
WarriorCase = namedtuple("WarriorCase", ["attack", "magic", "enemy_hero", "expected_effect"])

cases = [
    WarriorCase(attack="normal", magic=False, enemy_hero=False, expected_effect="effective"),
    WarriorCase(attack="normal", magic=False, enemy_hero=True, expected_effect="ineffective"),
    WarriorCase(attack="normal", magic=True, enemy_hero=False, expected_effect="effective"),
    WarriorCase(attack="normal", magic=True, enemy_hero=True, expected_effect="effective"),
    WarriorCase(attack="special", magic=False, enemy_hero=False, expected_effect="effective"),
    WarriorCase(attack="special", magic=False, enemy_hero=True, expected_effect="ineffective"),
    WarriorCase(attack="special", magic=True, enemy_hero=False, expected_effect="super_effective"),
    WarriorCase(attack="special", magic=True, enemy_hero=True, expected_effect="effective"),
]

print(f"Testing the {attack_effect.__name__} function:")
for case in cases:
    print(f"  With: attack={case.attack!r}, magic={case.magic!r}, "\
          f"enemy_hero={case.enemy_hero!r}, it should "\
          f"return: {case.expected_effect!r}. Outcome", end="...")
    
    # run student's function
    effect = attack_effect(attack=case.attack, magic=case.magic, enemy_hero=case.enemy_hero)

    # test returned value is same as expecteed
    try:
        assert effect == case.expected_effect, TEST_FAIL_MSG
        print("OK")
    except AssertionError:
        print("FAIL")
        raise

# if all test cases passed, print final message
print(TEST_PASS_MSG)

```

<!-- grade_id: cell-ce78b2aa0e59c35d, points: 0, grade: true, solution: false -->
```python
# This cell is reserved for grading, please do not change it.

```

---

## Problem 5: Coffee recommender

Write a function called `recommend_coffee` that recommends one of four coffees, depending on the following parameters:
- `dark_roast` (Boolean) - do you like dark roast?
- `cream` (Boolean) - do you like cream with your coffee?
- `hazelnut` (Boolean) - do you like hazelnut flavors with your coffee?

Your function should make the following recommendations:
- Like dark roast, but like neither cream nor hazelnut &rarr; `"americano"`
- Like dark roast, and must have both cream and hazelnut &rarr; `"house_special"`
- Like dark roast, and cream, but not hazelnut &rarr; `"latte"`
- Anything else &rarr; `"house_drip"`

<!-- grade_id: cell-685afe775c9bf961, points: 0, grade: false, solution: true -->
```python
# function recommending coffee's
def recommend_coffee(dark_roast, cream, hazelnut):
    """
    this function serves as a Coffee recommender based on flavor likings 
    
    Parameters 
    dark_roast : bool
    cream : bool
    hazelnut: bool
    
    Returns
    return : recommendations of americano, house_special, latte, and house_drip based on likes
    """
    if dark_roast == True and cream == False and hazelnut == False:
        return "americano"
    elif dark_roast == True and cream == True and hazelnut == True:
        return "house_special"
    elif dark_roast == True and cream == True and hazelnut == False:
        return "latte"
    else:
        return "house_drip"
```

```python
# Test your function here (and/or add more code blocks if you need to)
# You should probably play around with these variables to test your function's behavior
dark_roast = True
cream = False
hazelnut = True
coffee = recommend_coffee(dark_roast, cream, hazelnut)
print(coffee)
```

Output:
```
house_drip
```

When you are done testing your function, run the autograder below to check if your solution is correct.

The autograder will test several possible combinations of inputs at once.

```python
CoffeeCase = namedtuple("CoffeeCase", ["dark_roast", "cream", "hazelnut", "expected_coffee"])

cases = [
    CoffeeCase(dark_roast=False, cream=False, hazelnut=False, expected_coffee="house_drip"),
    CoffeeCase(dark_roast=False, cream=False, hazelnut=True, expected_coffee="house_drip"),
    CoffeeCase(dark_roast=False, cream=True, hazelnut=False, expected_coffee="house_drip"),
    CoffeeCase(dark_roast=False, cream=True, hazelnut=True, expected_coffee="house_drip"),
    CoffeeCase(dark_roast=True, cream=False, hazelnut=False, expected_coffee="americano"),
    CoffeeCase(dark_roast=True, cream=False, hazelnut=True, expected_coffee="house_drip"),
    CoffeeCase(dark_roast=True, cream=True, hazelnut=False, expected_coffee="latte"),
    CoffeeCase(dark_roast=True, cream=True, hazelnut=True, expected_coffee="house_special"),
]

print(f"Testing the {recommend_coffee.__name__} function:")
for case in cases:
    print(f"  With: dark_roast={case.dark_roast!r}, cream={case.cream!r}, "\
          f"hazelnut={case.hazelnut!r}, it should "\
          f"return: {case.expected_coffee!r}. Outcome", end="...")
    
    # run student's function
    coffee = recommend_coffee(dark_roast=case.dark_roast, cream=case.cream, 
                              hazelnut=case.hazelnut)

    # test returned value is same as expecteed
    try:
        assert coffee == case.expected_coffee, TEST_FAIL_MSG
        print("OK")
    except AssertionError:
        print("FAIL")
        raise

# if all test cases passed, print final message
print(TEST_PASS_MSG)

```

Output:
```
Testing the recommend_coffee function:
  With: dark_roast=False, cream=False, hazelnut=False, it should return: 'house_drip'. Outcome...OK
  With: dark_roast=False, cream=False, hazelnut=True, it should return: 'house_drip'. Outcome...OK
  With: dark_roast=False, cream=True, hazelnut=False, it should return: 'house_drip'. Outcome...OK
  With: dark_roast=False, cream=True, hazelnut=True, it should return: 'house_drip'. Outcome...OK
  With: dark_roast=True, cream=False, hazelnut=False, it should return: 'americano'. Outcome...OK
  With: dark_roast=True, cream=False, hazelnut=True, it should return: 'house_drip'. Outcome...OK
  With: dark_roast=True, cream=True, hazelnut=False, it should return: 'latte'. Outcome...OK
  With: dark_roast=True, cream=True, hazelnut=True, it should return: 'house_special'. Outcome...OK
Correct, great job! 👍
```

<!-- grade_id: cell-183d5dc8e4c7101a, points: 0, grade: true, solution: false -->
```python
# This cell is reserved for grading, please do not change it.

```